<a href="https://colab.research.google.com/github/toche7/DataAnalytic/blob/main/Lab5_BasicStat.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Lab 5: Basic Statistic



## Basic Statistic with NumPy



### การบวกค่าใน Array

ตัวอย่างสั้น ๆ ให้พิจารณาการคำนวณผลรวมของค่าทั้งหมดในอาร์เรย์
Python สามารถทำได้โดยใช้ฟังก์ชัน `sum`:


```python
import numpy as np
```

```python
L = np.random.random(100)
sum(L)
```

สำหรับการ `sum` ของ NumPy จะมีการใช้งานคล้ายกันดังนี้


```python
np.sum(L)
```

อย่างไรก็ตาม sum ของ NumPy จะได้ผลการคำนวณได้เร็วขึ้นมาก


```python
big_array = np.random.rand(1000000)
%timeit sum(big_array)
%timeit np.sum(big_array)
```

โปรดสังเกตุ: ฟังก์ชัน `sum` และฟังก์ชัน  `np.sum` จะไม่เหมือนกันซึ่งบางครั้งอาจทำให้เกิดความสับสนได้

โดยเฉพาะค่า option มีความหมายที่แตกต่างกันและ `np.sum` คำนึงถึงมิติอาร์เรย์หลายแบบดังที่เราจะเห็นในส่วนต่อไป



### Minimum and Maximum

ในทำนองเดียวกัน Python มีฟังก์ชัน `min` และ `max` ภายในซึ่งใช้เพื่อค้นหาค่าต่ำสุดและค่าสูงสุดของอาร์เรย์ที่กำหนด:


```python
min(big_array), max(big_array) # ลองสังเกตุด้วยว่าคำตอนนี้มีการคืนค่าออกมาเป็น tuple
```

```python
a = min(big_array), max(big_array)
```

```python
type(a)
```

ฟังก์ชันที่เกี่ยวข้องของ NumPy มีไวยากรณ์ที่คล้ายกันและทำงานได้เร็วขึ้นเช่นกัน

```python
np.min(big_array), np.max(big_array)
```

สำหรับ `min`, `max`, `sum` และการรวม NumPy อื่น ๆ การเรียกใช้คำสั่งสามารถทำได้อีกหนึ่งวิธีคือการสั่งผ่าน method ของ object ดังตัวอย่างดังต่อไปนี้

```python
 big_array.min(), big_array.max(), big_array.sum()
```

### Multi dimensional aggregates

Aggregate คือการจัดการข้อมูลที่มีอยู่หลายมิติให้เกิดเป็นข้อมูลใหม่เช่น ด้วยการบวก การค่าเฉลี่ย  หรือ จะเป็นการหาค่าสูงสุดต่ำสุด เพื่อนำไปใช้ประโยชน์ต่อไป

อย่างเช่นเรามีข้อมูลดังต่อไปนี้


```python
M = np.random.random((3, 4))
print(M)
```

โดยค่าเริ่มต้นฟังก์ชันการรวม NumPy แต่ละฟังก์ชันจะส่งคืนการรวมในอาร์เรย์ทั้งหมด:

```python
M.sum()
```

เราสามารถทำ aggregate เพื่อสร้างข้อมูลใหม่ที่เกิดจากการ  sum  โดยฟังก์ชัน sum จะใช้อาร์กิวเมนต์เพิ่มเติมโดยระบุ \*axis \* ซึ่งคำนวณการรวม ตัวอย่างเช่นเราสามารถค้นหาค่าต่ำสุดในแต่ละคอลัมน์ได้โดยระบุ `axis = 0`:


<img src=".Lab5_BasicStat.ipynb.upload/Picture1.png"   width="598.9891359999999px"  height="307px"  style="object-fit:cover"/>



```python
M.min(axis=0)
```

ฟังก์ชันจะส่งคืนค่าสี่ค่าซึ่งสอดคล้องกับสี่คอลัมน์ของตัวเลข

ในทำนองเดียวกันเราสามารถค้นหาค่าสูงสุดในแต่ละแถว:

```python
M.max(axis=1)
```

```python
M.sum(axis =0)
```

```python
M.sum(axis =1)
```

### Other aggregation functions

NumPy มีฟังก์ชันการรวมอื่น ๆ อีกมากมาย แต่เราจะไม่พูดถึงรายละเอียดที่นี่


**ฟังก์ชันสถิติเบื้องต้น:**

1. `np.mean()`: คำนวณค่าเฉลี่ย (average) ของสมาชิกในอาร์เรย์
2. `np.median()`: คำนวณค่ามัธยฐาน (median) ของสมาชิกในอาร์เรย์
3. `np.sum()`: คำนวณผลรวมของสมาชิกในอาร์เรย์
4. `np.std()`: คำนวณค่าส่วนเบี่ยงมาตรฐาน (standard deviation) ของสมาชิกในอาร์เรย์
5. `np.var()`: คำนวณค่าความแปรปรวน (variance) ของสมาชิกในอาร์เรย์
6. `np.min()`: ค้นหาค่าน้อยที่สุดในอาร์เรย์
7. `np.max()`: ค้นหาค่ามากที่สุดในอาร์เรย์
8. `np.percentile()`: คำนวณควอนไทล์ (percentiles) ที่ระบุของสมาชิกในอาร์เรย์
9. `np.histogram()`: คำนวณฮิสโตแกรมของชุดข้อมูล
10. `np.cov()`: คำนวณเมทริกซ์ความสหสัมพันธ์ (covariance) ของอาร์เรย์สองอาร์เรย์
11. `np.corrcoef()`: คำนวณเมทริกซ์ค่าสหสัมพันธ์ (correlation coefficient) ของอาร์เรย์สองอาร์เรย์
12. `np.histogram2d()`: คำนวณฮิสโตแกรมแบบ 2 มิติ

**การสร้างเลขสุ่ม:**

13. `np.random.rand()`: สร้างเลขสุ่มจากการกระจายเดียว (uniform distribution) [0, 1]
14. `np.random.randn()`: สร้างเลขสุ่มจากการกระจายมาตรฐาน (standard normal distribution)
15. `np.random.randint()`: สร้างเลขจำนวนเต็มสุ่มในช่วงที่ระบุ
16. `np.random.choice()`: สร้างตัวอย่างสุ่มจากอาร์เรย์




```python
M.std(axis = 0)
```

```python
np.cov(M)
```

```python
np.histogram(M)
```

## Example: What is the Average Height of US Presidents?



ฟังก์ชัน aggregate ที่มีอยู่ใน NumPy จะมีประโยชน์อย่างมากสำหรับการสรุปค่าของข้อมูลที่มี

ต่อไปนี้เป็นตัวอย่างง่าย ๆ เพื่อวิเคราะห์ความสูงของประธานาธิบดีสหรัฐฯทั้งหมด ข้อมูลนี้มีอยู่ในไฟล์ president\_heights.csv


```python
!head -10 president_heights.csv
```

```python
import pandas as pd
data = pd.read_csv('https://github.com/toche7/DataSets/raw/main/president_heights.csv')
heights = np.array(data['height(cm)'])
print(heights)
```

ตอนนี้เรามีอาร์เรย์ข้อมูลนี้แล้วเราสามารถคำนวณสถิติสรุปต่างๆได้:

```python
print("Mean height:       ", heights.mean())
print("Standard deviation:", heights.std())
print("Minimum height:    ", heights.min())
print("Maximum height:    ", heights.max())
```

จากการคำนวณข้างต้นจะได้ผลลัพธ์ออกมาค่าเดียว หากเราต้องการข้อมูลการกระจายอื่นๆ ต้องมีการคำนวณเพิ่มดังนี้


```python
print("25th percentile:   ", np.percentile(heights, 25))
print("Median:            ", np.median(heights))
print("75th percentile:   ", np.percentile(heights, 75))
```

เราเห็นว่าความสูงเฉลี่ยของประธานาธิบดีสหรัฐฯอยู่ที่ 182 ซม. หรือสูงแค่หกฟุต

เราจะมาลอง plot graph เพื่อที่จะได้เข้าใจข้อมูลมากยิ่งขึ้น


```python
import matplotlib.pyplot as plt
```

```python
plt.hist(heights)
plt.title('Height Distribution of US Presidents')
plt.xlabel('height (cm)')
plt.ylabel('number');
```

## Basic Statistic with Pandas

แม้ใน NumPy จะมีฟังก์ชันรองรับการคำนวณค่าสถิติต่าง ๆ อยู่บ้างแล้ว แต่เมื่อจะเทียบกับ Pandas นั้นเราจะพบว่า Pandas มีเครื่องมือหรือฟังก์ชันที่ทำขึ้นมาเพื่อการวิเคราะห์อย่างมากมาย

เมทอดและฟังก์ชัน Pandas ที่เกี่ยวข้องกับสถิติและการวิเคราะห์ข้อมูล โดยแบ่งกลุ่มตามหมวดหมู่ พร้อมกับชื่อเมทอดและฟังก์ชันในรูปแบบ 'pd.method' และคำอธิบายในภาษาไทย

**สถิติเบื้องต้น (Descriptive Statistics):**

1. `pd.mean()`: คำนวณค่าเฉลี่ย (average) ของสมาชิกใน DataFrame หรือ Series
2. `pd.median()`: คำนวณค่ามัธยฐาน (median) ของสมาชิกใน DataFrame หรือ Series
3. `pd.sum()`: คำนวณผลรวมของสมาชิกใน DataFrame หรือ Series
4. `pd.std()`: คำนวณค่าส่วนเบี่ยงมาตรฐาน (standard deviation) ของสมาชิกใน DataFrame หรือ Series
5. `pd.var()`: คำนวณค่าความแปรปรวน (variance) ของสมาชิกใน DataFrame หรือ Series
6. `pd.min()`: ค้นหาค่าน้อยที่สุดใน DataFrame หรือ Series
7. `pd.max()`: ค้นหาค่ามากที่สุดใน DataFrame หรือ Series
8. `pd.quantile()`: คำนวณควอนไทล์ (quantiles) ของข้อมูลใน DataFrame หรือ Series
9. `pd.describe()`: สร้างสรุปสถิติรวมรายละเอียด เช่น จำนวน, เฉลี่ย, ส่วนเบี่ยงมาตรฐาน, ค่าน้อยที่สุด, ควอนไทล์ 25%, 50%, 75%, และค่ามากที่สุดใน DataFrame หรือ Series
10. `pd.value_counts()`: นับค่าที่ไม่ซ้ำกันใน Series
11. `pd.mode()`: ค้นหาโหมด (ค่าที่เกิดขึ้นบ่อยที่สุด) ใน Series
12. `pd.cov()`: คำนวณเมทริกซ์สหสัมพันธ์ (covariance)
13. `pd.corr()`: คำนวณเมทริกซ์ค่าสหสัมพันธ์ (correlation)
14. `pd.rank()`: คำนวณอันดับของสมาชิกใน Series

**การรวมข้อมูลและการจัดกลุ่ม (Aggregation and Grouping):**

15. `pd.groupby()`: จัดกลุ่มข้อมูลตามหนึ่งหรือหลายคอลัมน์และคำนวณค่าสรุป
16. `pd.agg()`: ใช้ฟังก์ชันการรวมค่าหลาย ๆ ฟังก์ชันในข้อมูลที่ถูกจัดกลุ่ม
17. `pd.pivot_table()`: สร้างตารางหมุน (pivot table)
18. `pd.crosstab()`: คำนวณตารางตรวจสอบง่าย ๆ ของตัวแปรสองตัวหรือมากกว่า
19. `pd.resample()`: Resample ข้อมูลแบบ time-series
20. `pd.rolling()`: สร้างมุมมอง rolling ของข้อมูล time-series และคำนวณสถิติ rolling

**การทดสอบสมมติฐาน (Hypothesis Testing):**

21. `pd.ttest_ind()`: ทดสอบสมมติฐานสองตัวอย่างอิสระกันด้วย t-test
22. `pd.ttest_rel()`: ทดสอบสมมติฐานสองตัวอย่างที่เป็นคู่กันด้วย t-test
23. `pd.chi2_contingency()`: ทดสอบสมมติฐานความเชื่อมโยงในตารางการจัดหมวดหมู่ (contingency table) ด้วยทดสอบ chi-squared

**การเปลี่ยนแปลงข้อมูล (Data Transformation):**

24. `pd.cut()`: แบ่งข้อมูลต่อเขตต่อช่วงของตัวแปรต่อเนื่องเป็นช่วงๆ
25. `pd.qcut()`: แบ่งข้อมูลต่อควอนไทล์ (quantiles) ของตัวแปรต่อเนื่องเป็นกลุ่มๆ
26. `pd.shift()`: ย้ายข้อมูลไปข้างหน้าหรือข้างหลังในเวลาหรือแกน

**การสร้างแบบจำลองทางสถิติ (Statistical Modeling):**

27. `pd.OLS()`: หาแบบจำลองเชิงเส้นด้วยวิธีมาตรฐานจาก Ordinary Least Squares
28. `pd.logit()`: หาแบบจำลองโลจิสติกด้วยการจัดกลุ่มโมเดล
29. `pd.GEE()`: หาแบบจำลองเชิงเส้นที่ใช้ Generalized Estimating Equations (GEE) models

เหล่านี้เป็นเมทอดและฟังก์ชัน Pandas ที่ให้ความสามารถในการวิเคราะห์ทางสถิติ การรวมข้อมูล การเปลี่ยนแปลงข้อมูล การทดสอบสมมติฐาน และการสร้างแบบจำลอง เป็นเครื่องมือที่สำคัญสำหรับการจัดการและการวิเคราะห์ข้อมูลใน Python



```python
import pandas as pd
df = pd.read_csv('https://github.com/toche7/DataSets/raw/main/president_heights.csv')
```

```python
df
```

```python
df.max()
```

```python
df.min()
```

```python
df['height(cm)'].min()
```

```python
#dir(df)
```

```python
df.columns
```

```python
df['height(cm)'].mean()
```

```python
df['height(cm)'].median()
```

```python
df['height(cm)'].std()
```

```python
df['height(cm)'].var()
```

```python
df['height(cm)'].quantile([0.25,0.5,0.75])
```

```python
#help(df.quantile)
```

```python
df.describe()
```

```python
df['height(cm)'].value_counts()
```

```python
df['height(cm)'].value_counts()
```

### Ex. 1

จงปรับปรุงฟังก์ชันต่อไปนี้เพื่อคืนคำนวณหาค่า mean, median ของข้อมูล

เช่น

4, 8, 6, 5, 3, 2, 8, 8, 2, 10, 11

คำตอบที่คาดหวัง

mean = 6.090909090909091

median = 6.0



<!-- grade_id: ee22d1, points: 0, grade: false, solution: true -->
```python
import numpy as np
# Given data
data = np.array([4, 8, 6, 5, 3, 2, 8, 8, 2, 10,11])

def mm(data):
    mean = None # You Code here
    median = None # You Code here
    ### BEGIN SOLUTION
    mean = np.mean(data)
    median = np.median(data)
    ### END SOLUTION
    return mean, median


output = mm(data)
print('mean =', output[0])
print('median =', output[1])

# Expeted Result:
# mean = 6.090909090909091
# median = 6.0
```

```python
%pip install nose
```

<!-- grade_id: 26229d, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
data = np.array([4, 8, 6, 5, 3, 2, 8, 8, 2, 10])
output = mm(data)
assert_equal(output, (5.6,5.5))
```

### Ex. 2

จงปรับปรุงฟังก์ชันต่อไปนี้เพื่อคืนคำนวณหาค่า  Pearson correlation coefficient  และค่า covariance  ของข้อมูลสองชุด

เช่น

10, 20, 30, 40, 50

15, 25, 35, 45, 55

ค่าคำตอบที่คาดหวัง

Pearson Correlation Coeff: 1.0

Covariance: 250.0


<!-- grade_id: fc69ab, points: 0, grade: false, solution: true -->
```python
# Given data
d1 = np.array([10, 20, 30, 40, 50])
d2 = np.array([15, 25, 35, 45, 55])

def ccfunc(data1,data2):
    correlation_matrix =  None #You Code here
    covariance_matrix =   None #You Code here
    ### BEGIN SOLUTION
    correlation_matrix = np.corrcoef(data1, data2)
    covariance_matrix = np.cov(data1, data2)
    ### END SOLUTION
    return correlation_matrix[0][1],covariance_matrix[0][1]

print('Pearson Correlation Coeff:', ccfunc(d1,d2)[0])
print('Covariance:', ccfunc(d1,d2)[1])

```

<!-- grade_id: a8e087, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
data1 = np.array([10, 20, 30, 40, 50, 60])
data2 = np.array([15, 25, 35, 45, 55, 30])
from nose.tools import assert_equal, assert_almost_equal
assert_almost_equal(ccfunc(data1,data2), (0.6546536707079771,175.0),5)

```

### Ex. 3

จาก Ex. 2. ให้ใช้วิธีการคำนวณด้วย Pandas เพื่อคำนวณหาค่า Correlation และค่า Covariance



<!-- grade_id: 8350d2, points: 0, grade: false, solution: true -->
```python
# Given data
data = {
    'data1': [10, 20, 30, 40, 50],
    'data2': [15, 25, 35, 45, 55]
}
import pandas as pd


def ccfuncPanda(data):
    df = pd.DataFrame(data)
    correlation_matrix =  None #You Code here
    covariance_matrix =   None #You Code here
    ### BEGIN SOLUTION
    correlation_matrix = df.corr()
    covariance_matrix = df.cov()
    ### END SOLUTION
    return correlation_matrix.loc['data1', 'data2'], covariance_matrix.loc['data1', 'data2']

print('Pearson Correlation Coeff:', ccfuncPanda(data)[0])
print('Covariance:', ccfuncPanda(data)[1])

```

<!-- grade_id: 9e8d3b, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:

data = {
    'data1': [10, 20, 30, 40, 50, 60],
    'data2': [15, 25, 35, 45, 55, 30]
}

from nose.tools import assert_almost_equal
assert_almost_equal(ccfuncPanda(data)[0], 0.6546536707079771, 5)
assert_almost_equal(ccfuncPanda(data)[1], 175, 5)
```

```python

```